# Mastermind
**CS1302 Introduction to Computer Programming**
___

```python
%reset -f
import random
```

In this notebook, you will write a game called [*Mastermind*](https://en.wikipedia.org/wiki/Mastermind_(board_game)). Play the video below to learn about the rule of the game.

```python
%%html
<iframe width="912" height="513" src="https://www.youtube.com/embed/wsYPsrzCKiA" frameborder="0" allow="accelerometer; autoplay; encrypted-media; gyroscope; picture-in-picture" allowfullscreen></iframe>
```

## Code Generation (Demo)

The following function generates a random code, which is 
- a string of length `code_length` 
- with characters randomly chosen from those of the string `colors` with replacement.

```python
def get_code(colors, code_length):
    code = ''
    for i in range(code_length):
        code += random.choice(colors)
    return code
```

We can run the following to generate a code of `4` code pegs randomly with colors chosen from `'R'`ed, `'O'`range, `'Y'`ellow, `'G'`reen, `'B'`lue, and `'P'`urple. 

```python
get_code('ROYGBP', 4)
```

## Guess Validation

**Exercise** Given `colors` and `code_length`, check if `guess` is a valid code, i.e., a string of length `code_length` with characters from those of `colors`.

*Hint:* Solution template:
```Python
def valid_code(colors, code_length, guess):
    if len(guess) __ code_length:
        is_valid = ____
    else:
        for peg in code:
            for color in colors:
                if peg == color: ____
            else:
                is_valid = _____
                ____
        else:
            is_valid = ____
    return is_valid
```

<!-- grade_id: valid_code, points: 0, grade: false, solution: true -->
```python
def valid_code(colors, code_length, guess):
    # YOUR CODE HERE
    raise NotImplementedError()
    return is_valid
```

<!-- grade_id: test-valid_code, points: 1, grade: true, solution: false -->
```python
# tests
assert valid_code('RBG',1,'R') == True
assert valid_code('RBG',2,'B') == False
assert valid_code('RBG',2,'RP') == False
assert valid_code('RBG',0,'') == True
```

## Feedback Generation

**Exercise** Given `code` and `guess`, assign to `key` a feedback string, which starts with `black_key_pegs_count` characters of `'b'` (black key pegs) followed by `white_key_pegs_count` characters of `'w'` (white key pegs). 

*Hint:* Solution template:
```Python
def get_feedback(code, guess):
    black_key_pegs_count = white_key_pegs_count = counted = 0
    for i in range(len(guess)):
        if ___________________: 
            black_key_pegs_count += 1
            counted += 2**i
    for i in range(len(guess)):
        for j in range(len(code)):
            if  ______________________________ and not counted // 2**j % 2:
                white_key_pegs_count += 1
                counted += 2**j
                break
    key = 'b' * black_key_pegs_count + 'w' * white_key_pegs_count
    return key
```

<!-- grade_id: get_feedback, points: 0, grade: false, solution: true -->
```python
def get_feedback(code, guess):
    black_key_pegs_count = white_key_pegs_count = counted = 0
    # YOUR CODE HERE
    raise NotImplementedError()
    key = 'b' * black_key_pegs_count + 'w' * white_key_pegs_count
    return key
```

<!-- grade_id: test-get_feedback, points: 1, grade: true, solution: false -->
```python
# tests
def test_get_feedback(feedback,code,guess):
    feedback_ = get_feedback(code,guess)
    correct = feedback == feedback_
    if not correct:
        print(
            f'With code="{code}" and guess="{guess}", feedback should be "{feedback}", not "{feedback_}".'
        )
    assert correct

test_get_feedback(10*'b'+'w'*0,"RGBRGBRGBY","RGBRGBRGBY")
test_get_feedback(0*'b'+'w'*10,"RGBRGBRGBY","YRGBRGBRGB")
test_get_feedback(8*'b'+'w'*0,"RGRGRGRG","RGRGRGRG")
test_get_feedback(0*'b'+'w'*8,"RGRGRGRG","GRGRGRGR")
test_get_feedback(0*'b'+'w'*6,"RRRRGGG","GGGGRRR")
test_get_feedback(1*'b'+'w'*6,"RRRRGGG","GGGRRRR")
test_get_feedback(5*'b'+'w'*2,"RRRRGGG","RRRGGGR")
test_get_feedback(1*'b'+'w'*0,"RRRRGGG","RYYPPBB")
test_get_feedback(0*'b'+'w'*1,"RRRRG","GBBBB")
test_get_feedback(0*'b'+'w'*0,"RRRRG","YBBBB")
```

## Play the Game

After finishing the previous exercises, you can play the game as a code breaker against a random mastermind.

```python
# mastermind
max_num_guesses = 10
colors = 'ROYGBP'
code_length = 4

code = get_code(colors,code_length)
for i in range(max_num_guesses):
    while True:
        guess = input('Guess #{}: '.format(i))        
        if valid_code(colors,code_length,guess):
            break
        else:
            print('Invalid guess.')
    feedback = get_feedback(code,guess)
    print('Key:',feedback)
    if feedback == 'b'* code_length:
        print('You won!')
        break
else:
    print('You lose...')
```